---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Predictive Mean Matching (PMM)

```yaml
type: VideoExercise
key: 9fe4bb9735
xp: 50
```

`@projector_key`
dfc39e5c1fefef0cc571dfaba89a05af
